# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

Output:
```
'This is a docstring!'
```

$\beta=5$

```python
alpha = 4
```

```python
%who
print(datasets)

```

Output:
```
SVC	 X	 a	 alpha	 b	 c	 classifier	 datasets	 iris	 
predictions	 sklearn	 var_list	 y	 
<module 'sklearn.datasets' from '/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/datasets/__init__.py'>
```

```python
%who
```

Output:
```
SVC	 X	 a	 alpha	 b	 c	 classifier	 datasets	 iris	 
predictions	 sklearn	 var_list	 y
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['SVC',
 'X',
 'a',
 'alpha',
 'b',
 'c',
 'classifier',
 'datasets',
 'iris',
 'predictions',
 'sklearn',
 'var_list',
 'y']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris datase
# YOUR CODE HERE
#raise NotImplementedError()
raise?
```

Output:
```
Object `raise` not found.
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf", gamma='auto')
classifier.fit(X, y)
predictions = classifier.predict(X)

print(classifier, "\n and I am just testing")
#print?
print(train_test_split)
#train_test_split?
X_train, X_test, y_train, y_test = train_test_split(
    X, y, test_size=0.33, random_state=42)
```

Output:
```
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto', kernel='rbf',
    max_iter=-1, probability=False, random_state=None, shrinking=True,
    tol=0.001, verbose=False) 
 and I am just testing
<function train_test_split at 0x7fb346f44268>
```

```python
ord('b')
ord?

%ls *
a, b, c = 1, 2, 3

print(a)
print(b)
print(c)

import sklearn
sklearn.__version__

y_train[99]
```

Output:
```
index.rst  Week1-assignment.ipynb
1
2
3
2
[0;31mSignature:[0m [0mord[0m[0;34m([0m[0mc[0m[0;34m,[0m [0;34m/[0m[0;34m)[0m[0;34m[0m[0;34m[0m[0m
[0;31mDocstring:[0m Return the Unicode code point for a one-character string.
[0;31mType:[0m      builtin_function_or_method
```

```python
%who
```

```python
from sklearn.model_selection import train_test_split
```

```python
train_test_split?
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC()

X.shape

classifier.fit(X, y)

predictions = classifier.predict(X)
```

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Train test split (or cross-validation)
4. Fit
5. Predict

```python
predictions
```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
x = 0;
if x > 0:
    raise NotImplementedError()

```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```